Negative affect influences the computations underlying food choice in bulimia nervosa

Author response:

The following is the authors’ response to the original reviews

eLife Assessment

This study provides a valuable contribution to understanding how negative affect influences food-choice decision making in bulimia nervosa, using a mechanistic approach with a drift diffusion model (DDM) to examine the weighting of tastiness and healthiness attributes. The solid evidence is supported by a robust crossover design and rigorous statistical methods, although concerns about low trial counts, possible overfitting, and the absence of temporally aligned binge-eating measures limit the strength of causal claims. Addressing modeling transparency, sample size limitations, and the specificity of mood induction effects, would enhance the study's impact and generalizability to broader populations.

We thank the Editor and Reviewers for their summary of the strengths of our study, and for their thoughtful review and feedback on our manuscript. We apologize for the confusion in how we described the multiple steps performed to ensure that the hierarchical model reported in the main text was the best fit for the data but was not overfitted. Regarding “model transparency,” as described in our response to Reviewer 1 below, we have now more clearly explained (with references) that the use of hierarchical estimation procedures allows for information sharing across participants, which improves the reliability and stability of parameter estimates—even when the number of trials per individual is small. We have clarified for the less familiar reader how our Bayesian model selection criterion penalizes models with more parameters (e.g., more complex models).

Details about model diagnostics, recoverability, and posterior predictive checks are all provided in the Supplementary Materials. We have clarified how these steps ensure that the parameters we estimate are identifiable and interpretable, while confirming that the model can reproduce key patterns in the data, ultimately supporting the validity of the winning model. Additionally, we have provided all scripts for estimating the models by linking to our public Github repository. Furthermore, we have edited language throughout to eliminate any implication of causal claims and acknowledged the limitation of the small sample size. Given these efforts, we are concerned that the current wording about “modeling transparency” in the public eLife Assessment may inadvertently misrepresent the modeling practices in our paper. Would it be possible to revise or remove that particular phrase to better reflect the steps we have taken? We believe this would help avoid confusion for readers.

We have also taken additional steps to ensure that we have used “appropriate and validated methodology in line with current state-of-the-art," and we have added references to recent papers supporting our approaches.

All changes in the revised text are marked in blue.

Public Reviews:

Reviewer #1 (Public review):

Summary:

Using a computational modeling approach based on the drift diffusion model (DDM) introduced by Ratcliff and McKoon in 2008, the article by Shevlin and colleagues investigates whether there are differences between neutral and negative emotional states in:

(1) The timings of the integration in food choices of the perceived healthiness and tastiness of food options between individuals with bulimia nervosa (BN) and healthy participants.

(2) The weighting of the perceived healthiness and tastiness of these options.

Strengths:

By looking at the mechanistic part of the decision process, the approach has the potential to improve the understanding of pathological food choices. The article is based on secondary research data.

Weaknesses:

I have two major concerns and a major improvement point.

The major concerns deal with the reliability of the results of the DDM (first two sections of the Results, pages 6 and 7), which are central to the manuscript, and the consistency of the results with regards to the identification of mechanisms related to binge eating in BN patients (i.e. last section of the results, page 7).

(1) Ratcliff and McKoon in 2008 used tasks involving around 1000 trials per participant. The Chen et al. experiment the authors refer to involves around 400 trials per participant. On the other hand, Shevlin and colleagues ask each participant to make two sets of 42 choices with two times fewer participants than in the Chen et al. experiment. Shevlin and colleagues also fit a DDM with additional parameters (e.g. a drift rate that varies according to subjective rating of the options) as compared to the initial version of Ratcliff and McKoon. With regards to the number of parameters estimated in the DDM within each group of participants and each emotional condition, the 5- to 10-fold ratio in the number of trials between the Shevlin and colleagues' experiment and the experiments they refer to (Ratcliff and McKoon, 2008; Chen et al. 2022) raises serious concerns about a potential overfitting of the data by the DDM. This point is not highlighted in the Discussion. Robustness and sensitivity analyses are critical in this case.

We thank the Reviewer for their thoughtful critique. We agree that a limited number of trials can impede reliable estimation, which we acknowledge in the Discussion section. However, we used a hierarchical estimation approach which leverages group information to constrain individual-level estimates. This use of group-level parameters to inform individual-level estimates reduces overfitting and noise that can arise when trial counts are low, and the regularization inherent in hierarchical fitting prevents extreme parameter estimates that could arise from noisy or limited data (Rouder & Lu, 2005). As a result, hierarchical estimation has been repeatedly shown to work well in settings with low trial counts, including as few as 40 trials per condition (Lerche et al., 2017; Ratcliff & Childers, 2015; Wiecki et al., 2013). In addition, previous applications of the time-varying DDM to food choice task data has included experiments with as few as 60 trials per condition (Maier et al., 2020). We have added references to these more recent approaches and specifically note their advantages for the modeling of tasks with fewer trials. Finally, our successful parameter recovery described in the Supplementary Materials supports the robustness of the estimation procedure and the reliability of our results.

The authors compare different DDMs to show that the DDM they used to report statistical results in the main text is the best according to the WAIC criterion. This may be viewed as a robustness analysis. However, the other DDM models (i.e. M0, M1, M2 in the supplementary materials) they used to make the comparison have fewer parameters to estimate than the one they used in the main text. Fits are usually expected to follow the rule that the more there are parameters to estimate in a model, the better it fits the data. Additionally, a quick plot of the data in supplementary table S12 (i.e. WAIC as a function of the number of parameters varying by food type in the model - i.e. 0 for M0, 2 for M1, 1 for M2 and 3 for M3) suggests that models M1 and potentially M2 may be also suitable: there is a break in the improvement of WAIC between model M0 and the three other models. I would thus suggest checking how the results reported in the main text differ when using models M1 and M2 instead of M3 (for the taste and health weights when comparing M3 with M1, for τS when comparing M3 with M2). If the differences are important, the results currently reported in the main text are not very reliable.

We thank the Reviewer for highlighting that it would be helpful to explicitly note that we specifically selected WAIC as one of two methods to assess model fit because it penalizes for model complexity. We now explicitly state that, in addition to being more robust than other metrics like AIC or BIC when comparing hierarchical Bayesian models like those in the current study, model fit metrics like WAIC penalize for model complexity based on the number of parameters (Watanabe, 2010). Therefore, more complex models (i.e., those with more parameters) do not automatically have lower WAIC. Additionally, we now more clearly note that our second method to assess model fit, posterior predictive checks, demonstrate that only model M3 can reproduce key behavioral patterns present in the empirical data. As described in the Supplementary Materials, M1 and M2 miss key patterns in the data. In summary, we used best practices to assess model fit and reliability (Wilson & Collins, 2019): results from the WAIC comparison (which penalizes models with more parameters) and results from posterior predictive checks align in showing that M3 provided the best fit to our data. We have added a sentence to the manuscript to state this explicitly.

(2) The second main concern deals with the association reported between the DDM parameters and binge eating episodes (i.e. last paragraph of the results section, page 7). The authors claim that the DDM parameters "predict" binge eating episodes (in the Abstract among other places) while the binge eating frequency does not seem to have been collected prospectively. Besides this methodological issue, the interpretation of this association is exaggerated: during the task, BN patients did not make binge-related food choices in the negative emotional state. Therefore, it is impossible to draw clear conclusions about binge eating, as other explanations seem equally plausible. For example, the results the authors report with the DDM may be a marker of a strategy of the patients to cope with food tastiness in order to make restrictive-like food choices. A comparison of the authors' results with restrictive AN patients would be of interest. Moreover, correlating results of a nearly instantaneous behavior (i.e. a couple of minutes to perform the task with the 42 food choices) with an observation made over several months (i.e. binge eating frequency collected over three months) is questionable: the negative emotional state of patients varies across the day without systematically leading patients to engage in a binge eating episode in such states.

I would suggest in such an experiment to collect the binge craving elicited by each food and the overall binge craving of patients immediately before and after the task. Correlating the DDM results with these ratings would provide more compelling results. Without these data, I would suggest removing the last paragraph of the Results.

We thank the Reviewer for these interesting and important suggestions, and we agree that claims about causal connections between our decision parameters and symptom severity metrics would be inappropriate. Per the Reviewer’s suggestions, we have eliminated the use of the word “predict” to describe the tested association with symptom metrics. We also agree that more time-locked associations with craving ratings and near-instantaneous behavior would be useful, and we have added this as an important direction for future research in the discussion. However, associating task-based behavior with validated self-report measures that assess symptom severity over long periods of time that precede the task visit (e.g., over the past 2 weeks in depression, over the past month in eating disorders) is common practice in computational psychiatry, psychiatric neuroimaging, and clinical cognitive neuroscience (Hauser et al., 2022; Huys et al., 2021; Wise et al., 2023), and this approach has been used several times specifically with food choice tasks (Dalton et al., 2020; Steinglass et al., 2015). We have revised the language throughout the manuscript to clarify: the results suggest that individuals whose task behavior is more reactive to negative affect tend to be the most symptomatic, but the results do not allow us to determine whether this reactivity causes the symptoms.

In response to this Reviewer’s important point about negative affect not always producing loss-of-control eating in individuals with BN, we now explicitly note that while several studies employing ecological momentary assessments (EMA) have repeatedly shown that increases in negative affect significantly increase the likelihood of subsequent loss-of-control eating (Alpers & Tuschen-Caffier, 2001; Berg et al., 2013; Haedt-Matt & Keel, 2011; Hilbert & Tuschen-Caffier, 2007; Smyth et al., 2007), not all loss-of-control eating occurs in the context of negative affect. We further note that future studies should integrate food choice task data pre and post-affect inductions with measures capturing the specific frequency of loss of control eating episodes that occur during states of high negative affect.

(3) My major improvement point is to tone down as much as possible any claim of a link with binge eating across the entire manuscript and to focus more on the restrictive behavior of BN patients in between binge eating episodes (see my second major concern about the methods). Additionally, since this article is a secondary research paper and since some of the authors have already used the task with AN patients, if possible I would run the same analyses with AN patients to test whether there are differences between AN (provided they were of the restrictive subtype) and BN.

We appreciate the Reviewer’s very helpful suggestions. We have adjusted our language linking loss-of-control eating frequency with decision parameters, and we have added sentences focusing on the implications for the restrictive behavior of patients with BN between binge eating episodes. In the Supplementary Materials, we have added an analysis of the restraint subscale of the EDE-Q and confirmed no relationship with parameters of interest. While we agree additional analyses with AN patients would be of interest, this is outside the scope of the paper. Our team have collected data from individuals with AN using this task, but not with any affect induction or measure of affect. Therefore, we have added this important direction for future research to the discussion.

Reviewer #2 (Public review):

Summary:

Binge eating is often preceded by heightened negative affect, but the specific processes underlying this link are not well understood. The purpose of this manuscript was to examine whether affect state (neutral or negative mood) impacts food choice decision-making processes that may increase the likelihood of binge eating in individuals with bulimia nervosa (BN). The researchers used a randomized crossover design in women with BN (n=25) and controls (n=21), in which participants underwent a negative or neutral mood induction prior to completing a food-choice task. The researchers found that despite no differences in food choices in the negative and neutral conditions, women with BN demonstrated a stronger bias toward considering the 'tastiness' before the 'healthiness' of the food after the negative mood induction.

Strengths:

The topic is important and clinically relevant and methods are sound. The use of computational modeling to understand nuances in decision-making processes and how that might relate to eating disorder symptom severity is a strength of the study.

Weaknesses:

The sample size was relatively small and may have been underpowered to find differences in outcomes (i.e., food choice behaviors). Participants were all women with BN, which limits the generalizability of findings to the larger population of individuals who engage in binge eating. It is likely that the negative affect manipulation was weak and may not have been potent enough to change behavior. Moreover, it is unclear how long the negative affect persisted during the actual task. It is possible that any increases in negative affect would have dissipated by the time participants were engaged in the decision-making task.

We thank the Reviewer for their comments on the strengths of the paper, and for highlighting these important considerations regarding the sample demographics and the negative affect induction. As in the original paper that focused only on ultimate food choice behaviors, we now specifically acknowledge that the study was only powered to detect small to medium group differences in the effect of negative emotion on these final choice behaviors.

Regarding the sample demographics, we agree that the study’s inclusion of only female participants is a limitation. Although the original decision for this sampling strategy was informed by data suggesting that bulimia nervosa is roughly six times more prevalent among females than males (Udo & Grilo, 2018), we now note in the discussion that our female-only sample limits the generalizability of the findings.

We also agree with the Reviewer’s noted limitations of the negative mood induction, and based on the reviewer’s suggestions, we have expanded our original description of these limitations in the Discussion. Specifically, we now note that although the task was completed immediately after the affect induction, the study did not include intermittent mood assessments throughout the choice task, so it is unclear how long the negative affect persisted during the actual task.

Reviewer #3 (Public review):

Summary:

The study uses the food choice task, a well-established method in eating disorder research, particularly in anorexia nervosa. However, it introduces a novel analytical approach - the diffusion decision model - to deconstruct food choices and assess the influence of negative affect on how and when tastiness and healthiness are considered in decision-making among individuals with bulimia nervosa and healthy controls.

Strengths:

The introduction provides a comprehensive review of the literature, and the study design appears robust. It incorporates separate sessions for neutral and negative affect conditions and counterbalances tastiness and healthiness ratings. The statistical methods are rigorous, employing multiple testing corrections.

A key finding - that negative affect induction biases individuals with bulimia nervosa toward prioritizing tastiness over healthiness - offers an intriguing perspective on how negative affect may drive binge eating behaviors.

Weaknesses:

A notable limitation is the absence of a sample size calculation, which, combined with the relatively small sample, may have contributed to null findings. Additionally, while the affect induction method is validated, it is less effective than alternatives such as image or film-based stimuli (Dana et al., 2020), potentially influencing the results.

We agree that the limited sample size and specific affect induction method may have contributed to the null model-agnostic behavioral findings. Based on this Reviewer’s and Reviewer 2’s comments, we have added these factors to our acknowledgements of limitations in the discussion.

Another concern is the lack of clarity regarding which specific negative emotions were elicited. This is crucial, as research suggests that certain emotions, such as guilt, are more strongly linked to binge eating than others. Furthermore, recent studies indicate that negative affect can lead to both restriction and binge eating, depending on factors like negative urgency and craving (Leenaerts et al., 2023; Wonderlich et al., 2024). The study does not address this, though it could explain why, despite the observed bias toward tastiness, negative affect did not significantly impact food choices.

We thank the Reviewer for raising these important points and possibilities. In the Supplementary Materials, we have added an additional analysis of the specific POMS subscales that comprise the total negative affect calculation that was reported in the original paper (Gianini et al., 2019). We also report total negative affect scores from the POMS in the main text. Ultimately, we found that, across both groups, the negative affect induction increased responses related to anger, confusion, depression, and tension while reducing vigor.

We agree with the Reviewer that factors like negative urgency and cravings are relevant here. The study did not collect any measures of craving, and in response to Reviewer 1 and this Reviewer, we now note in the discussion that replication studies including momentary craving assessments will be important. While we do not have any measurements of cravings, we did measure negative urgency. The original paper (Gianini et al., 2019) did not find that negative urgency was related to restrictive food choices. We have now repeated those analyses, and we also were unable to find any meaningful patterns related to negative urgency. Nonetheless, we have added an analysis of negative urgency scores and decision parameters to the Supplementary Materials.

Recommendations for the authors:

Reviewer #1 (Recommendations for the authors):

Please improve the description of the computational methods: the fit of the DDM, the difference between the models used in the DDM, and the difference between the DDM model and the models used in the linear mixed models (the word "model" is at the end confusing as it may refer either to the DDM or to the statistical analysis of the DDM parameters).

We thank the Reviewer for highlighting the unclear language. We have updated the main text to clarify when the term “model” refers to the DDM itself versus the regression models assessing DDM parameters. As described above, we have clarified that both tests of model fit (WAIC and posterior predictive checks) suggest that Model 3 was the best fit to the data. We have also clarified the differences between the tested models in the Supplementary Materials.

Please avoid reporting estimates of main effects in statistical models when an interaction is included: the estimates of the main effects may be heavily biased by the interaction term (this can be checked by re-running the model without the interaction term).

We sincerely appreciate the Reviewer’s comment regarding the interpretation of main effects in the presence of significant interaction terms. In the revised manuscript, we no longer discuss significant main effects and instead focus on interpreting the interaction terms.

Additionally, to help unpack interaction effects, we now include exploratory simple effects analyses in the supplementary materials. Simple effects analyses allow us to examine the effects of one independent variable at specific values of other independent variables (Aiken et al., 1991; Brambor et al., 2006; Jaccard & Turrisi, 2003; Winer et al., 1991).

Supplementary tables S5 and S6 are excessive: there is no third-level interaction (supplementary tables S3 and S4) to justify a split between BN and healthy participants. Please perform rather a descending regression. Accordingly, the results reported in the second paragraph of page 7 should be entirely rewritten.

We agree with the Reviewer’s suggestion that these tables are unnecessary. We have updated them to include details about simple effects analyses described above. We have revised the main text to reflect these changes.

The words such as "predictive" indicating a causality link is used in several places in the manuscript including the supplementary materials while the experimental design does not allow such claims. This should be rephrased.

We agree with the Reviewer that the term “predicted” in the main text improperly suggested a causal relationship between symptom severity and DDM parameters that our methods cannot evaluate. We have updated the main text with more appropriate language. However, our use of the term “predicted” in the Supplementary Materials refers to predicting the probability of a choice based on trial-level features which is standard use of the term in the computational cognitive modeling literature (Piray et al., 2019; Wilson & Collins, 2019; Zhang et al., 2020).

The word "evaluated" appears twice in line 42 of the supplementary materials. Same with "in" at line 50.

Thank you very much for highlighting this. We have removed the repeated words.

Reviewer #2 (Recommendations for the authors):

(1) I think it would be helpful if the authors noted in the Methods how long the food-choice task took. Prior research has suggested that in-lab mood inductions are very short-lasting (e.g., max 7 minutes) and it is likely that the task itself may have impacted the mood states of participants. Expanding on this in the Discussion/limitations seems important.

The Reviewer raises an important point regarding the duration of our affect manipulation. Since we did not measure mood during or after the Food Choice Task, we cannot determine how long these effects persisted. We have added this limitation to the discussion section, noting that the absence of continuous affect measures following mood induction is a widespread limitation in the field.

(2) Personally, I was a bit confused about what data the researchers were using to extrapolate information on whether or not participants were considering healthiness or tastiness. How was this operationalized? Is this an assumption being made based on how quickly someone chose a low-fat vs. high-fat food?

We thank this Reviewer for highlighting that our models’ complexity warrants a more thorough explanation.

Since we collected tastiness and healthiness attribute ratings during the first phase of the Food Choice Task, we can use those values to determine how these attribute values influence decision-making. Independently, foods were classified as low-fat or high-fat based on their objective properties (i.e., the percentage of calories from fat). However, the primary information we used to compute model parameters were participants’ attribute ratings, choices, and response times.

In these models, the drift rate parameter captures the speed and direction of evidence accumulation. As the unsigned magnitude of the drift rate increases, the decision-maker is making up their mind more quickly. Once the evidence accumulates to a response boundary, the option associated with that boundary is selected. A positive drift rate means they are moving toward choosing one option (i.e., upper boundary), and a negative drift rate means they are moving toward choosing the other (i.e., lower boundary). In these decisions, decision-makers often consider multiple attributes, such as perceived healthiness and tastiness. Each of these attributes can influence the evidence accumulation process with different strengths, or weights.

In addition, decision-makers do not consider all attributes at the same time. Inspired by earlier work on multi-attribute decision-making (Maier et al., 2020; Sullivan & Huettel, 2021), our modeling approach computes a parameter (i.e., relative attribute onset) which captures the time delay between when each attribute starts influencing the evidence accumulation process. This parameter gives us a way to estimate when decision-makers are considering different attributes, and tells us how much influence each attribute has, because if the attribute starts late, it has less time to influence the decision. These models use a piecewise drift rate function to describe how evidence changes over time within a trial: sometimes the decision maker only considers taste, sometimes only health, and other times both. Importantly, models with a relative attribute onset parameter can produce key behavioral patterns observed in mouse-tracking studies that models without this parameter are unable to replicate (Maier et al., 2020).

In summary, the computational model describes decision-makers’ behaviors (what they would choose, and how fast they would choose) using different potential values of the drift weights and relative start time parameters. We then used Bayesian estimation methods to compare the model's predictions to the actual data. By examining how reaction times and choices change depending on the attribute values of the presented options, the model allows us to infer when each attribute is considered, and how strongly it influences the final choice.

We have clarified this in the main text.

Reviewer #3 (Recommendations for the authors):

I wonder whether there were any measures concerning negative affect before and after the mood induction? This would make it clearer whether there was a significant change before and after. If different emotions were assessed, which emotion showed the strongest change?

We thank the Reviewer for flagging this point. We realize that the main text did not make it clear that mood was assessed before and after the mood induction using the POMS (McNair et al., 1989). While these analyses were conducted and the results were reported in the original manuscript (Gianini et al., 2019), we now report them in the main text for completeness. Additionally, we added more details about how specific emotions changed by analyzing the subscales of the POMS in the Supplementary Materials. As mentioned above, we found that, across both groups, the negative affect induction increased responses related to anger, confusion, depression, and tension while reducing vigor.

Thank you again for your consideration and for the reviewers’ comments and suggestions. We believe their incorporation has significantly strengthened the paper. In addition, thank you for the opportunity to publish our work in eLife. We look forward to hearing your response.

References

Aiken, L. S., West, S. G., & Reno, R. R. (1991). Multiple regression: Testing and interpreting interactions. Sage Publications, Inc.

Alpers, G. W., & Tuschen-Caffier, B. (2001). Negative feelings and the desire to eat in bulimia nervosa. Eating Behaviors, 2(4), 339–352. https://doi.org/10.1016/S1471-0153(01)00040-X

Berg, K. C., Crosby, R. D., Cao, L., Peterson, C. B., Engel, S. G., Mitchell, J. E., & Wonderlich, S. A. (2013). Facets of negative affect prior to and following binge-only, purge-only, and binge/purge events in women with bulimia nervosa. Journal of Abnormal Psychology, 122(1), 111–118. https://doi.org/10.1037/a0029703

Brambor, T., Clark, W. R., & Golder, M. (2006). Understanding Interaction Models: Improving Empirical Analyses. Political Analysis, 14(1), 63–82. https://doi.org/10.1093/pan/mpi014

Dalton, B., Foerde, K., Bartholdy, S., McClelland, J., Kekic, M., Grycuk, L., Campbell, I. C., Schmidt, U., & Steinglass, J. E. (2020). The effect of repetitive transcranial magnetic stimulation on food choice-related self-control in patients with severe, enduring anorexia nervosa. International Journal of Eating Disorders, 53(8), 1326–1336. https://doi.org/10.1002/eat.23267

Gianini, L., Foerde, K., Walsh, B. T., Riegel, M., Broft, A., & Steinglass, J. E. (2019). Negative affect, dietary restriction, and food choice in bulimia nervosa. Eating Behaviors, 33, 49–54. https://doi.org/10.1016/j.eatbeh.2019.03.003

Haedt-Matt, A. A., & Keel, P. K. (2011). Revisiting the affect regulation model of binge eating: A meta-analysis of studies using ecological momentary assessment. Psychological Bulletin, 137(4), 660–681. https://doi.org/10.1037/a0023660

Hauser, T. U., Skvortsova, V., Choudhury, M. D., & Koutsouleris, N. (2022). The promise of a model-based psychiatry: Building computational models of mental ill health. The Lancet Digital Health, 4(11), e816–e828. https://doi.org/10.1016/S2589-7500(22)00152-2

Hilbert, A., & Tuschen-Caffier, B. (2007). Maintenance of binge eating through negative mood: A naturalistic comparison of binge eating disorder and bulimia nervosa. International Journal of Eating Disorders, 40(6), 521–530. https://doi.org/10.1002/eat.20401

Huys, Q. J. M., Browning, M., Paulus, M. P., & Frank, M. J. (2021). Advances in the computational understanding of mental illness. Neuropsychopharmacology, 46(1), 3–19. https://doi.org/10.1038/s41386-020-0746-4

Jaccard, J., & Turrisi, R. (2003). Interaction effects in multiple regression (2nd ed.). Sage Publications, Inc.

Lerche, V., Voss, A., & Nagler, M. (2017). How many trials are required for parameter estimation in diffusion modeling? A comparison of different optimization criteria. Behavior Research Methods, 49(2), 513–537. https://doi.org/10.3758/s13428-016-0740-2

Maier, S. U., Raja Beharelle, A., Polanía, R., Ruff, C. C., & Hare, T. A. (2020). Dissociable mechanisms govern when and how strongly reward attributes affect decisions. Nature Human Behaviour, 4(9), Article 9. https://doi.org/10.1038/s41562-020-0893-y

McNair, D., Lorr, M., & Droppleman, L. (1989). Profile of mood states (POMS).

Piray, P., Dezfouli, A., Heskes, T., Frank, M. J., & Daw, N. D. (2019). Hierarchical Bayesian inference for concurrent model fitting and comparison for group studies. PLOS Computational Biology, 15(6), e1007043. https://doi.org/10.1371/journal.pcbi.1007043

Ratcliff, R., & Childers, R. (2015). Individual differences and fitting methods for the two-choice diffusion model of decision making. Decision, 2(4), 237–279. https://doi.org/10.1037/dec0000030

Rouder, J. N., & Lu, J. (2005). An introduction to Bayesian hierarchical models with an application in the theory of signal detection. Psychonomic Bulletin & Review, 12(4), 573–604. https://doi.org/10.3758/BF03196750

Smyth, J. M., Wonderlich, S. A., Heron, K. E., Sliwinski, M. J., Crosby, R. D., Mitchell, J. E., & Engel, S. G. (2007). Daily and momentary mood and stress are associated with binge eating and vomiting in bulimia nervosa patients in the natural environment. Journal of Consulting and Clinical Psychology, 75(4), 629–638. https://doi.org/10.1037/0022-006X.75.4.629

Steinglass, J., Foerde, K., Kostro, K., Shohamy, D., & Walsh, B. T. (2015). Restrictive food intake as a choice—A paradigm for study. International Journal of Eating Disorders, 48(1), 59–66. https://doi.org/10.1002/eat.22345

Sullivan, N., & Huettel, S. A. (2021). Healthful choices depend on the latency and rate of information accumulation. Nature Human Behaviour, 5(12), Article 12. https://doi.org/10.1038/s41562-021-01154-0

Udo, T., & Grilo, C. M. (2018). Prevalence and Correlates of DSM-5–Defined Eating Disorders in a Nationally Representative Sample of U.S. Adults. Biological Psychiatry, 84(5), 345–354. https://doi.org/10.1016/j.biopsych.2018.03.014

Watanabe, S. (2010). Asymptotic Equivalence of Bayes Cross Validation and Widely Applicable Information Criterion in Singular Learning Theory. Journal of Machine Learning Research, 11, 3571–3594.

Wiecki, T. V., Sofer, I., & Frank, M. J. (2013). HDDM: Hierarchical Bayesian estimation of the drift-diffusion model in Python. Frontiers in Neuroinformatics, 7. https://doi.org/10.3389/fninf.2013.00014

Wilson, R. C., & Collins, A. G. (2019). Ten simple rules for the computational modeling of behavioral data. eLife, 8, e49547. https://doi.org/10.7554/eLife.49547

Winer, B. J., Brown, D. R., & Michels, K. M. (1991). Statistical principles in experimental design (3rd ed). McGraw-Hill.

Wise, T., Robinson, O. J., & Gillan, C. M. (2023). Identifying Transdiagnostic Mechanisms in Mental Health Using Computational Factor Modeling. Biological Psychiatry, 93(8), 690–703. https://doi.org/10.1016/j.biopsych.2022.09.034

Zhang, L., Lengersdorff, L., Mikus, N., Gläscher, J., & Lamm, C. (2020). Using reinforcement learning models in social neuroscience: Frameworks, pitfalls and suggestions of best practices. Social Cognitive and Affective Neuroscience, 15(6), 695–707. https://doi.org/10.1093/scan/nsaa089

Changes from the earlier version to the revised version

Abstract

Removed in the revised version: Abstract

Earlier version: Individuals often consume tasty, calorically dense foods in response to negative emotions, a phenomenon exemplified by notions of “stress eating” and “comfort food.” While this link between food and mood can become pathological in binge eating, the decision-making processes underlying this link are poorly understood. Here, we investigated the impact of acute increases in negative affect on when and how strongly the perceived tastiness and healthiness of foods influence food choices in healthy adults and individuals with bulimia nervosa (BN), an eating disorder characterized by cycles of over- and underconsumption of food. In a randomized crossover design, 25 women with BN and 21 healthy controls completed two sessions where they received either a neutral or negative affect induction and then completed a food choice task. Using a time-varying diffusion decision model, we assessed how negative affect influences food choice dynamics for high- and low-fat foods. In the neutral affect condition, individuals with BN considered tastiness relative to healthiness of high-fat foods sooner than healthy controls but maintained a restrictive food choice policy by reducing the weight on tastiness. After a negative affect induction, both groups showed a stronger bias towards considering tastiness before healthiness, but this bias was exaggerated in individuals with BN. This affect-induced bias for high-fat foods predicted more frequent subjective binge episodes over three months. These results provide insights into how negative emotion influences food choices and may explain why binge eating in BN is more likely during high negative affect, while dietary restriction is more likely during low negative affect.

Revised version: Individuals often consume tasty, calorically dense foods in response to negative emotions, a phenomenon exemplified by notions of “stress eating” and “comfort food.” While this link between food and mood can become pathological in binge eating, the decision-making processes underlying this link are poorly understood. Here, we investigated the impact of acute increases in negative affect on when and how strongly the perceived tastiness and healthiness of foods influence food choices in healthy adults and individuals with bulimia nervosa (BN), an eating disorder characterized by cycles of over- and under-consumption of food. In a randomized crossover design, 25 women with BN and 21 healthy controls completed two sessions where they received either a neutral or negative affect induction and then completed a food choice task. Using a time-varying diffusion decision model, we assessed how negative affect influences food choice dynamics for high- and low-fat foods. In the neutral affect condition, individuals with BN considered tastiness relative to healthiness of high-fat foods sooner than healthy controls but maintained a restrictive food choice policy by reducing the weight on tastiness. After a negative affect induction, both groups showed a stronger bias towards considering tastiness before healthiness, but this bias was exaggerated in individuals with BN. This affect-induced bias for high-fat foods predicted more frequent subjective binge episodes over three months. These results provide insights into how negative emotion influences food choices and may explain why binge eating in BN is more likely during high negative affect, while dietary restriction is more likely during low negative affect.

Introduction

Every day, multiple times a day, people must make decisions about what to eat. A growing body of research suggests that these decisions involve considerations of different attributes of the available food options (e.g., healthiness and tastiness), and that these attributes often present the decision-maker with challenging conflicts (e.g. “Should I choose the healthier, less delicious food or the better tasting, less healthy food?”). Difficulties resolving these conflicts can take on pathological forms in individuals with eating disorders, and some eating disorders are characterized by food choices that are biased toward one extreme (e.g., toward subjectively less tasty, but healthier, low-fat foods in anorexia nervosa; (Foerde et al., 2015)). However, food choices seem to episodically oscillate between extremes in the case of bulimia nervosa (BN), which is commonly marked by prolonged periods of rigid dietary restriction punctuated by out-of-control binge eating of highly palatable foods. Although subclinical forms of this oscillation are ubiquitous in examples of “yo-yo dieting” and “dietary lapses,” surprisingly little cognitive neuroscience research to date has attempted to explain it.

Earlier version: Self-report data suggest that rising negative affect is one reliable predictor of the switch from more to less restrictive food choices (Konttinen et al., 2010). Although some individuals reduce intake in response to stress (Hill et al., 2022), research using both animal models and healthy adult samples show that stress can precipitate the overconsumption of highly palatable foods (Dionysopoulou et al., 2021; Jacques et al., 2019; Tomiyama, 2019). This connection between negative emotion and food choice is even more pronounced in binge eating. Inducing negative affect in the laboratory environment disproportionately increases food intake among individuals with binge eating (Telch & Agras, 1996). Ecological momentary assessment data show that binge eating in the natural environment is also much more likely to occur in the context of increasing and unstable negative emotions (Alpers & Tuschen-Caffier, 2001; Berg et al., 2013; Haedt-Matt & Keel, 2011; Hilbert & Tuschen-Caffier, 2007; Smyth et al., 2007). In addition, reactivity to negative emotions, or negative urgency, has been shown to predict more frequent binge eating in individuals with BN (Fischer et al., 2013; Schnepper et al., 2021).

Revised version: Self-report data suggest that rising negative affect is one reliable predictor of the switch from more to less restrictive food choices (Konttinen et al., 2010). Although some individuals reduce intake in response to stress (Hill et al., 2022), research using both animal models and healthy adult samples show that stress can precipitate the overconsumption of highly palatable foods (Dionysopoulou et al., 2021; Jacques et al., 2019; Tomiyama, 2019). This connection between negative emotion and food choice is even more pronounced in binge eating. Inducing negative affect in the laboratory environment disproportionately increases food intake among individuals with binge eating (Telch & Agras, 1996). Ecological momentary assessment data show that binge eating in the natural environment is also much more likely to occur in the context of increasing and unstable negative emotions (Alpers & Tuschen-Caffier, 2001; Berg et al., 2013; Haedt-Matt & Keel, 2011; Hilbert & Tuschen-Caffier, 2007; Smyth et al., 2007). In addition, reactivity to negative emotions, or negative urgency, has been shown to predict or correlate with more frequent binge eating in individuals with BN (Fischer et al., 2013; Schnepper et al., 2021).

Although momentary increases in negative emotions are closely tied to the overconsumption of highly palatable foods, the cognitive processes underlying this tight link, particularly in the case of binge eating, are not well understood. Studies to date have failed to find support for the notion that increased negative affect, specifically stress, impairs inhibitory control processes in healthy adults (Allen et al., 2022) or in BN (Dreyfuss et al., 2017; Westwater et al., 2021). While another study found that healthy adults and individuals with BN differed in how stress impacted visual processing-related neural activity, they failed to find support for the hypothesis that increased negative affect abnormally increased limbic neural responses to palatable food stimuli in BN (Collins et al., 2017).

Earlier version: We recently tested whether negative affect potentiates binge eating via its impact on decision-making, specifically by promoting more high-fat, tasty food choices in BN (Gianini et al., 2019). Individuals with BN and healthy controls completed a food choice task after a negative and neutral affect induction (Figure 1). Counter to expectations, negative affect did not influence ultimate food choices in either group, and women with BN showed restrictive food choices whether neutral or negative affect was induced: they chose low-fat foods more often (Figure 2A), and these choices were strongly predicted by healthiness ratings (Figure 2B). However, these analyses focused primarily on individuals’ ultimate choices and may have overlooked the possibility that even small changes in negative affect could influence more subtle aspects of the food decision-making process leading up to those choices.

Revised version: We recently tested possibility: that negative affect abnormally impacts decision-making in BN by promoting food choices that include the high-fat, appetitive foods that are typically consumed during binge-eating episodes (Gianini et al., 2019). Individuals with BN and healthy controls completed a food choice task after a negative and neutral affect induction (Figure 1). Counter to expectations, negative affect did not influence ultimate food choices in either group, and women with BN showed restrictive food choices whether neutral or negative affect was induced: they chose low-fat foods more often (Figure 2A), and these choices were strongly predicted by healthiness ratings (Figure 2B). However, these analyses focused primarily on individuals’ ultimate choices and may have overlooked the possibility that even small changes in negative affect could influence more subtle aspects of the food decision-making process leading up to those choices.

Affect induction and task design.

(A) Study timeline. BN and HC participants completed the study tasks on two separate days. On session 1, participants were randomly assigned to the neutral mood or negative mood induction before completing the Food Choice Task. On session 2, participants experienced the alternative mood induction before completing another run of the Food Choice Task. The mood inductions involved combinations of music and autobiographical writing. (B) Food Choice Task. During the Food Choice Task, both BN and HC participants rated 43 food items across three phases. In the Tastiness Ratings and Healthiness phases, participants rated each item on a 5-point Likert scale from Bad to Good and Unhealthy to Healthy, respectively. In the Choice phase, participants indicated their strength of preference for a presented food item, compared to a personally tailored neutral reference item.

Removed in the revised version: (A) Food choice task behavior estimated from regression models.

Earlier version: Re-analysis of the raw data excluding outlier response times (2.5% of trials n = 85) replicated original findings (Gianini et al., 2019). While both groups were less likely to choose high-fat foods (over the neutral reference item) than low-fat foods (over the neutral reference item), the BN group was even less likely than the HC group to choose high-fat food items. However, we did not identify any significant effects of the Affect Condition on choices. Error bars indicate 95% confidence intervals of the estimated effects. (B) Influence of health and taste ratings on food choice. Health ratings influenced food choice more in the BN group than in the HC group. Within the HC group, food choice was influenced more strongly by taste ratings than health ratings. Error bars indicate standard errors of the estimated coefficients. Note. Corresponding statistics are presented in Table S8. HC = healthy controls; BN = bulimia nervosa.

Revised version: (A) Food choice task behavior estimated from regression models. Re-analysis of the raw data excluding outlier response times (2.5% of trials n = 85) replicated original findings (Gianini et al., 2019). While both groups were less likely to choose high-fat foods (over the neutral reference item) than low-fat foods (over the neutral reference item), the BN group was even less likely than the HC group to choose high-fat food items. However, we did not identify any significant effects of the Affect Condition on choices. Error bars indicate 95% confidence intervals of the estimated effects. (B) Influence of health and taste ratings on food choice. Health ratings influenced food choice more in the BN group than in the HC group. Within the HC group, food choice was influenced more strongly by taste ratings than health ratings. Error bars indicate standard errors of the estimated coefficients. Note. Corresponding statistics are presented in Table S8. HC = healthy controls; BN = bulimia nervosa.

Added in the revised version: These more subtle features can be quantified with computational cognitive modeling. Sequential sampling models leverage both choice and response time data to estimate parameters that capture individual differences in decision-making dynamics (Forstmann et al., 2016; Ratcliff & Smith, 2004). These models assume that decision-making involves the noisy accumulation of evidence over time until a predetermined threshold is reached, leading to a response. Therefore, both ultimate responses and response times are meaningful signals about the cognitive processes underlying decision-making. Though originally developed to describe perceptual, psychomotor, and memory tasks, these models can be applied to value-based decision-making tasks by incorporating participants’ ratings of different attributes for the stimuli presented in a trial. This allows researchers to examine how and when these attributes influence the decision-making process (Busemeyer et al., 2019; Clithero, 2018).

Earlier version: These more subtle features can be quantified with cognitive computational modeling. Research indicates that food choices are influenced not only by the weights assigned to attributes like healthiness and tastiness, but also by the time at which these attributes enter the decision-making process (Chen et al., 2022; HajiHosseini & Hutcherson, 2021; Hutcherson & Tusche, 2022; Maier et al., 2020; Sullivan et al., 2015; Sullivan & Huettel, 2021). Using time-varying sequential sampling models, studies in healthy individuals have demonstrated that decision-makers assign different subjective weights to healthiness and taste attribute ratings of foods (i.e., they have differing degrees to which they influence the evidence accumulation rate); that healthiness and taste attribute information can enter the decision process at different times; and that these relative weighting strengths and timing of attribute consideration have distinct influences on ultimate food choices (Maier et al., 2020). This notion is consistent with theories of both emotion- and food-craving regulation, which posit that there are separate attention deployment and stimulus valuation steps involved in these regulation processes (e.g., (Giuliani & Berkman, 2015; Gross, 2015; Han et al., 2018).

Revised version: Research using sequential sampling models indicates that food choices are influenced not only by the weights assigned to attributes like healthiness and tastiness, but also by the time at which these attributes enter the decision-making process (Chen et al., 2022; HajiHosseini & Hutcherson, 2021; Hutcherson & Tusche, 2022; Maier et al., 2020; Sullivan et al., 2015; Sullivan & Huettel, 2021). Specifically, using time-varying sequential sampling models, studies have demonstrated that healthy individuals assign different subjective weights to healthiness and taste attribute ratings of foods (i.e., they have differing degrees to which they influence the evidence accumulation rate); that healthiness and taste attribute information can enter the decision process at different times; and that these relative weighting strengths and timing of attribute consideration have distinct influences on ultimate food choices (Maier et al., 2020). This notion is consistent with theories of both emotion- and food-craving regulation, which posit that there are separate attention deployment and stimulus valuation steps involved in these regulation processes (e.g., (Giuliani & Berkman, 2015; Gross, 2015; Han et al., 2018).

Notably, even if one of these decision parameters (e.g., weights on healthiness or tastiness, attribute onsets) biased individuals toward selecting a low-fat (or high-fat) food, the other parameter could compensate for it to ultimately generate the opposite response (Figure 3). For example, if taste information was considered before healthiness information, but healthiness information was weighted strongly enough, a decision-maker could choose a healthier, instead of a tastier, food (Figure 3B). These latent aspects of the decision-making process, either orthogonally or in parallel, could be sensitive to state changes like increases in negative affect and could be altered in BN. Specifically, these latent decision parameters could help explain why dysregulated consumption of tasty, high-fat foods in BN is more likely during periods of high negative affect, but restricted intake of lower-fat foods is more likely during periods of low negative affect. In one case, increasing negative affect could alter an individual’s attribute weights, either reducing their weight on healthiness (dashed blue line) or increasing their weight on tastiness (dashed red line), which would increase their likelihood of choosing binge-type foods that are considered to be highly palatable yet also unhealthy (Figure 3C). Alternatively, increasing negative affect could cause individuals with BN to consider healthiness information later in the decision-making process, giving taste information more time to bias choices towards binge-type foods (Figure 3D). Clarifying these mechanisms is critical for translational efforts because data suggest that the weighting of attributes and the time at which they are considered have different neural substrates (Maier et al., 2020), and the treatment offered to patients should depend on which scenario more accurately reflects their pathology.

Example evidence accumulation trajectories predicted by the starting time diffusion decision model (stDDM).

Earlier version: Each trajectory is a simulated agent that considers one attribute alone before beginning to consider both tastiness and healthiness attributes together. (A) Tastiness onset delay. The dashed, green trajectory represents a case where positive healthiness information about a low-fat food item quickly biases the agent towards the accept threshold before taste information has time to influence the evidence accumulation process. The solid purple trajectory illustrates a case where evidence related to a high-fat food is initially biased towards the reject threshold while the aversive healthiness information dominates the evidence accumulation process. Once information about the item’s appetitive tastiness comes online, the evidence accumulation changes its trajectory towards the accept threshold. (B) Healthiness onset delay. The dashed, orange trajectory illustrates a case where evidence related to the highly appetitive tastiness attribute dominates the evidence accumulation process, influencing the trajectory to terminate at the accept threshold of a high-fat food before the healthiness attribute is considered. For the solid brown trajectory, a less appetitive high-fat food item is ultimately rejected once aversive healthiness information enters the evidence accumulation process. (C-D) Hypothesized model of affect-induced binge-eating. Each trajectory is a simulated agent with BN making a decision involving a high-fat food. During neutral affect (gray solid line), the high-fat food is trending towards the accept threshold until the onset of aversive healthiness information biases the trajectory towards the reject threshold. (C) Attribute-weight hypothesis. During negative affect, either the attribute weight for taste information increases (red dashed line) or the attribute weight for healthiness information decreases (blue dashed line). In both cases, the trajectory is ultimately biased towards the accept threshold. (D) Attributeonset hypothesis. During negative affect, the initial delay shifts and taste information is accumulated longer before healthiness information comes online. With a longer delay in the onset of healthiness information, the evidence accumulation for the high-fat food has enough time to reach the accept threshold.

Revised version: Each trajectory is a simulated agent that considers one attribute alone before beginning to consider both tastiness and healthiness attributes together. (A) Tastiness onset delay. The dashed, green trajectory represents a case where positive healthiness information about a low-fat food item quickly biases the agent towards the accept threshold before taste information has time to influence the evidence accumulation process. The solid purple trajectory illustrates a case where evidence related to a high-fat food is initially biased towards the reject threshold while the aversive healthiness information dominates the evidence accumulation process. Once information about the item’s appetitive tastiness comes online, the evidence accumulation changes its trajectory towards the accept threshold. (B) Healthiness onset delay. The dashed, orange trajectory illustrates a case where evidence related to the highly appetitive tastiness attribute dominates the evidence accumulation process, influencing the trajectory to terminate at the accept threshold of a high-fat food before the healthiness attribute is considered. For the solid brown trajectory, a less appetitive high-fat food item is ultimately rejected once aversive healthiness information enters the evidence accumulation process. (C-D) Hypothesized model of affect-induced binge-eating. Each trajectory is a simulated agent with BN making a decision involving a high-fat food. During neutral affect (gray solid line), the high-fat food is trending towards the accept threshold until the onset of aversive healthiness information biases the trajectory towards the reject threshold. (C) Attribute-weight hypothesis. During negative affect, either the attribute weight for taste information increases (red dashed line) or the attribute weight for healthiness information decreases (blue dashed line). In both cases, the trajectory is ultimately biased towards the accept threshold. (D) Attribute-onset hypothesis. During negative affect, the initial delay shifts and taste information is accumulated longer before healthiness information comes online. With a longer delay in the onset of healthiness information, the evidence accumulation for the high-fat food has enough time to reach the accept threshold.

Here, we applied a diffusion decision model (DDM) with a time-varying drift rate to food choice and response time data from our controlled, randomized crossover study in women with BN and healthy controls (Figure 1A). This model allowed us to investigate the potential exaggerated influence of negative affect on the dynamics of food-specific decision-making in BN. Specifically, we tested how food type (low-fat/high-fat) and affective state (neutral/negative) influenced decision-makers’ attribute weights and onsets.

Our hypotheses were derived from the clinical phenomenology underlying BN. As outlined above, individuals with BN seem to vacillate between unstable decision dynamics: many individuals with the disorder maintain a relatively restrictive diet until negative emotional states induce binge-eating episodes and subsequent compensatory behaviors (i.e., purging). Given this instability, we predicted that our BN participants’ food choices would be influenced by components of the decision-making process whose opposing dynamics could be disrupted by increasing negative affect.

We first hypothesized that in states of neutral affect and for high-fat foods, one of these opposing dynamics would compensate for the other to promote ultimate choices that are restrictive. This could be achieved by either a) weighing taste information more strongly but considering it later than healthy controls, or b) weighing tastiness less strongly but considering it earlier than healthy controls. Second, we hypothesized that, in BN, increasing negative affect would abnormally impact at least one of these latent aspects of decision-making to increase the bias towards high-fat foods. Even though prior analyses indicate that high-fat foods were not ultimately chosen more often, we predicted that after a negative affect induction, individuals would either a) be slower to consider healthiness information, but still weight taste information just as strongly, or b) they would put less weight on healthiness information, but still consider taste information sooner.

Because these combinations require one feature of the decision process to compensate for another, they would result in a much less stable dynamic compared to one in which information weights and onset times are concordant. For example, when healthiness information is considered sooner and is more strongly weighted than taste information, foods low in healthiness would almost always be rejected. In contrast, when healthiness information is considered later but is more strongly weighed than taste information, food choice behavior can be more easily perturbed with state-based changes in weights or the delay in attribute onsets. Here, we show evidence for these unstable decision dynamics that can account for a bias in BN toward low-fat, healthier foods in states of relatively low negative affect, but toward high-fat, tastier foods in states of high negative affect. Although negative affect delayed the onset of the consideration of healthiness information among both HC and BN participants, this effect was more pronounced in the BN group. Ultimately, we found that although individuals with BN consistently weighed healthiness information more strongly, negative affect delayed its entry into the evidence accumulation process. With this longer delay, taste information has more time to bias decision-makers towards high-fat food choices before healthiness information can come online.

Results

Negative affect aberrantly delays consideration of healthiness information among women with BN

Added in the revised version: We analyzed data from individuals with BN (n = 25) and healthy controls (HC, n = 21) who completed a computerized Food Choice Task following negative and neutral affect inductions, across two visits (Figure 1). Mood was assessed before and after the affect inductions using the Profile of Mood States (POMS; (McNair et al., 1989)). Negative affect scores were calculated by summing the five negative subscale scores (e.g., anger, confusion, depression, fatigue, and tension) and subtracting the positive subscale score (e.g., vigor). Negative affect scores increased significantly for both groups following the negative affect induction (Gianini et al., 2019, reproduced in Table S1).

Earlier version: We analyzed data from individuals with BN (n = 25) and healthy controls (HC, n = 21) who participated in a computerized Food Choice Task following negative and neutral affect inductions, across two visits (Figure 1). To quantify the specific computations performed during Food Choice Task decision-making, we used a time-varying diffusion decision model (DDM; (Ratcliff & McKoon, 2008)). This time-varying model introduces a starting time parameter that indicates the delay with which different attributes enter the evidence accumulation process (Chen et al., 2022; Lombardi & Hare, 2021; Maier et al., 2020). Using this starting time DDM (stDDM), we can determine how individuals weigh healthiness versus taste information and when those attributes influence food-related decision-making in BN and HC groups in each affective state. Parameter estimates were fitted to data that included individual-level response times, choices, z-scored healthiness and tastiness ratings, and affect induction. We specified a model where individual parameters were drawn from four separate hyperparameters that varied both by group (BN or HC) and by condition (neutral or negative affect induction). To test our hypotheses about the interactions between affect and food type (low-fat or high-fat) on food choice, we allowed each attribute weight parameter (ωTaste and ωHealth) and the relative-starting time parameter (τs) to vary as a function of food type. The drift rate determining the evidence update can be written as follows: v(t)={ωTaste⋅VDTasteIf τs<0Λ0<t<|τs|ωHealth⋅VDHealthIf τs<0Λ0<t<τsωTaste⋅VDTaste+ωHealth⋅VDHealthIf t>τs (1)

Revised version: To quantify the specific computations performed during Food Choice Task decision-making, we used a time-varying diffusion decision model (DDM; (Ratcliff & McKoon, 2008)). This time-varying model introduces a starting time parameter that indicates the delay with which different attributes enter the evidence accumulation process (Chen et al., 2022; Lombardi & Hare, 2021; Maier et al., 2020). Using this starting time DDM (stDDM), we can determine how individuals weigh healthiness versus taste information and when those attributes influence food-related decision-making in BN and HC groups in each affective state. The stDDM was fitted to data that included individual-level response times, choices, z-scored healthiness and tastiness ratings, and affect induction.

Added in the revised version: We specified a stDDM where individual parameters were drawn from four separate hyperparameters that varied both by group (BN or HC) and by condition (neutral or negative affect induction). To test our hypotheses about the interactions between affect and food type (low-fat or high-fat) on food choice, we allowed each attribute weight parameter (ωTaste and ωHealth) and the relative-starting time parameter (τS) to vary as a function of food type. The drift rate determining the evidence update can be written as follows: v(t)={ωTaste⋅VDTasteωHealth⋅VDHealthωTaste⋅VDTaste+ωHealth⋅VDHealthIf τS<0∧0<t<|τS|If τS>0∧0<t<τSIf >|τS| (1)

Earlier version: where ωTaste is the subjective weight given to tastiness, ωHealth, is the subjective weight given to healthiness, VDTaste and VDHeaith are the value differences in taste and healthiness attributes, respectively, and τs is the time at which the taste and healthiness attributes come into the evidence accumulation process. If τs > 0, healthiness information is accumulated first, and evidence from taste comes into consideration at time τs, whereas τs < 0 means that taste information is accumulated first and evidence from healthiness information starts to come into the decision process at time |τs|. To test our hypotheses, medians of the posterior distributions for subject-level parameters were used as the dependent variables in three separate, fully factorial (Group-by-Affect Condition-by Food Type), mixed-effects linear regressions, one for each parameter of interest. Each of these binary categorical variables were contrast coded using the values -1 and 1, allowing us to separately assess group differences in the neutral and negative Affect Conditions.

Revised version: where ωTaste is the subjective weight given to tastiness, ωHealth. is the subjective weight given to healthiness, VDTaste and VDHealth are the value differences in taste and healthiness attribute ratings, respectively, and τS is the time at which the taste and healthiness attributes come into the evidence accumulation process. If τS > 0, healthiness information is accumulated first, and evidence from taste comes into consideration at time τS, whereas ωS < 0 means that taste information is accumulated first and evidence from healthiness information starts to come into the decision process at time |τS|. To test our hypotheses, medians of the posterior distributions for subject-level parameters were used as the dependent variables in our statistical analyses. We ran three separate, fully factorial (Group-by-Affect Condition-by Food Type), mixed-effects linear regressions, one for each parameter of interest. Each of these binary categorical variables were contrast coded using the values −1 and 1, allowing us to separately assess group differences across condition.

Results exploring additional parameters (boundary separation (α), non-decision time (τND), and starting point bias (z)) are presented in the Supplementary Materials along with model comparisons, model identifiability, posterior predictive checks, and parameter recovery exercises.

Added in the revised version: First, we assessed how relative start time parameters varied across groups as a function of food type in each condition. We found a significant three-way interaction between group, condition, and food type (Table S2; β = −0.32, t = −2.22, p = 0.028), indicating that the two groups differed in how negative affect influenced their attribute bias patterns across food types. Visual inspection of the data suggested this interaction was driven by group differences across food types in the neutral condition that disappeared in the negative affect condition, and by a more pronounced influence of negative affect on attribute biases in BN (Figure 4A). Specifically, in the neutral condition, the HC group showed greater taste bias for low-fat versus high-fat foods (M = −0.430 vs. −0.286), while individuals with BN showed greater taste bias for high-fat versus low-fat foods (M = −0.293 vs. −0.450; Table S3). The negative affect induction eliminated this crossover pattern, increasing taste biases across both food types and groups, but with greater increases observed in the BN group (Tables S2, S3).

Removed in the revised version: First, we assessed how relative start time parameters varied across groups as a function of food type in each condition. A Group-by-Food Type interaction indicated that in the neutral affect condition, the BN group, relative to the HC group, focused for longer on the taste information of high-fat foods than the taste information of low-fat foods (Figure 4A, Table S1; β = −0.30, t = −2.97, p = 0.003). In other words, the BN group showed a stronger bias towards first considering the taste attribute of high-fat foods. A main effect of Affect Condition indicates that both groups showed an increased bias toward considering taste information sooner than healthiness information in a negative emotional state (Figure 4A; β = −0.20, t = −2.63, p = 0.009; Table S2). However, an interaction between Group and Affect Condition (β = −0.23, t = 2.27, p = 0.025) suggests that the bias toward earlier consideration of taste information after the negative affect induction was stronger among the BN group than in HC. Additionally, this negative affect-induced bias favoring taste information significantly reduced the group-dependent Food Type distinctions that were present in the neutral state, as indicated by a significant three-way interaction of Food Type, Group, and Affect Condition (β = 0.31, t = 2.19, p = 0.031).

Parameter estimates.

Colors indicate diagnosis: purple = bulimia nervosa (BN); green = healthy controls (HC). Error bars represent the standard error of the mean. (A) Attribute onset. In the neutral condition, we observed a Group by Food Type cross-over effect: while the BN group showed a greater initial bias towards accumulating tastiness information of high-fat foods than low-fat foods, the HC group showed a greater tastiness information bias for low-fat foods than high-fat foods. After the negative affect induction, any Food Type-based distinctions disappeared, and both groups’ biases towards tastiness information increased, but this effect was more pronounced in the BN group. (B) Weight on tastiness information. The HC group put more weight on tastiness information than the BN group. The negative affect induction reduced tastiness weights for the HC group, but not the BN group. (C) Weight on healthiness information. The BN group put more weight on healthiness information than the HC group, especially for high-fat foods. The negative affect induction did not have significant effects on healthiness weights for either group.

Reduced tastiness weights are abnormally unaffected by negative affect in BN

Earlier version: Our second and third models assessed how each attribute weight varied across groups as a function of food type in each condition (Figure 4B-C, Table S3-6). In the neutral condition, the BN group had a significantly reduced ωtaste compared to the HC group (Table S3; β = −0.37, t = −5.83, p = 2.78 × 10−8) but a significantly increased ωhealth compared to the HC group (Table S4, β = 0.37, t = 5.20, p = 5.26 × 10−7). While both groups put less weight on taste information of high-fat foods than low-fat foods (β = −0.23, t = −3.72, p = 2.92 × 10−4), only the BN group had negative weights on high-fat taste information. A significant Group-by-Food Type interaction for health attribute weights (Table S4) indicates that the BN group had an especially strong ωheaith for high-fat foods compared to healthy controls (β = 0.29, t = 2.97, p = 0.004). This reduced weight on tastiness and increased weight on healthiness information in the BN group may explain how, despite focusing for longer on taste information, individuals with BN ultimately made restrictive eating choices in states of low negative affect.

Revised version: Our second and third regression models assessed how each attribute weight varied across groups as a function of food type in each condition (Figure 4B-C). For taste weights, a significant two-way interaction between group and condition (Table S4; β = 0.18, t = 2.11, p = 0.036) indicated that although the HC group placed consistently higher weight on taste compared to the BN group, the HC group reduced their weighting of taste information following negative affect induction while the BN group maintained similar taste weights across conditions (Table S5). A main effect of food type (β = −0.23, t = − 3.69,p < 0.001) showed that both groups placed less weight on taste information for high-fat versus low- fat foods. In contrast, healthiness weights were insensitive to negative affect induction across both groups. However, a significant interaction between group and food type (Table S6; β = 0.29, t = 2.94, p = 0.004) indicated that the BN group placed consistently higher weights on health information, particularly for high-fat foods, while the HC group showed no difference in health weighting across food types (Table S7).

Removed in the revised version: A Group-by-Condition interaction suggested that while the weight on taste information tended to be lower in both groups after the negative affect induction (β = −0.17, t = −2.73, p = 0.007), this effect was less pronounced in the BN group (β = 0.17, t = 2.07, p = 0.040). Indeed, exploratory analyses within each group (Table S5) revealed that the HC showed significant negative affect-induced decreases in ωtaste (β = −0.17, t = −2.84, p = 0.006), while the BN group did not show significant affect-based changes in ωtaste (β = 0.01, t = 0.09, p = 0.929). This suggests that while individuals with BN generally had lower weights for taste information than healthy adults, negative affect only influenced how healthy adults weighted taste information. Neither group exhibited any affect-based changes in ωheaith (Tables S4, S6).

More severe symptoms of BN are linked to affect-induced delays in processing healthiness information

Earlier version: To assess the clinical significance of our findings, we sought to connect model parameters to retrospective self-reported symptom severity. For these exploratory analyses, we used median values of the individuallevel posterior distributions of the attribute onset parameter τs, the parameter which showed an exaggerated response to the affect induction in the BN group. These values were then separately added to models with Group and Food Type to predict the severity of two symptoms in the BN group: the frequency of objectively large binge eating episodes (OBEs) and subjectively large binge eating episodes (SBEs) in the past three months. We examined both types of out-of-control eating episodes because they both contribute to a diagnosis of BN according to ICD-11 (World Health Organization, 2022), and some data suggest that SBEs are more strongly related to negative affect (Brownstone et al., 2013; Brownstone & Bardone-Cone, 2021; Fitzsimmons-Craft et al., 2014). Indeed, we found that that greater negative affect-based decreases in τs for high-fat foods relative to low-fat foods were associated with more frequent SBE in the past three months (Table S7;β = −32.46, t = −3.87,p = 1.08 × 10−4). In other words, those who focused for longer on the tastiness of high-fat foods after the negative affect induction reported more SBE episodes in the past three months (Figure 5). Conversely, affect-based changes in Ts were unrelated to OBE frequency (Table S7).

Revised version: To assess the clinical significance of our findings, we sought to connect stDDM parameters to retrospective self-reported symptom severity. For these exploratory analyses, we used median values of the individual-level posterior distributions of the attribute onset parameter τs, the parameter which showed an exaggerated response to the affect induction in the BN group. These parameter values were then separately added to mixed-effects negative binomial regression models with Group and Food Type to assess their association with two symptom severity measures in the BN group: the frequency of objectively large binge eating episodes (OBEs) and subjectively large binge eating episodes (SBEs) in the past three months. We examined both types of out-of-control eating episodes because they both contribute to a diagnosis of BN according to ICD-11 (World Health Organization, 2022), and some data suggest that SBEs are more strongly related to negative affect (Brownstone et al., 2013; Brownstone & Bardone-Cone, 2021; Fitzsimmons-Craft et al., 2014). Indeed, we found that that greater negative affect-based decreases in τS for high-fat foods relative to low-fat foods were associated with more frequent SBE in the past three months (Table S8; β = −46.39, t = −2.80, p = 0.005). In other words, those who focused for longer on the tastiness of high-fat foods after the negative affect induction reported more SBE episodes in the past three months (Figure 5). Conversely, affect-based changes in τS were unrelated to OBE frequency (Table S8).

Affect-induced changes in information onset predict subjective binge episodes.

In the negative affect condition (right facet), longer delays in accumulating healthiness information (i.e., reduced τS) for high-fat foods compared to low-fat foods was associated with more frequent subjective binge episodes. Line type and shape refer to Food Type: Solid lines and circles = Low fat foods; dashed lines and triangles = High fat foods.

Discussion

Earlier version: Negative affect is a known precursor to over-eating and dietary rule-breaking (Frayn & Knauper, 2018). Healthy adults, and even rodents, can exhibit dysregulated food consumption and altered food choices after periods of stress or other negative emotions (Dionysopoulou et al., 2021; Jacques et al., 2019; Konttinen et al., 2010; Tomiyama, 2019). Negative affect’s influence on eating behavior is particularly pronounced among individuals with clinical diagnoses like BN, who tend to have large and out-of-control eating episodes in states of strong negative emotion (Cardi et al., 2015; Wonderlich et al., 2022). Prior studies have demonstrated that outside of binge-eating episodes, and in states of lower negative affect, individuals with BN typically engage in restrictive eating behaviors, consuming some low-fat foods and largely avoiding high-fat foods (Alpers & Tuschen-Caffier, 2004; Bjorlie et al., 2022; Elran-Barak et al., 2015; Kales, 1990). We previously found that women with BN demonstrated these restrictive food choices on a food-related decision task, regardless of their affective state (Gianini et al., 2019). However, by leveraging a computational model that captures the decision-making process underlying ultimate choices, we found that negative affect more strongly biased individuals with BN than healthy adults, increasing the amount of time that taste information alone influenced their decisions involving food. This affect-induced delay in the consideration of healthiness versus tastiness may account for the fact that negative affect is a common precursor to taste-driven overeating in healthy adults, and the fact that it is a particularly strong predictor of taste-driven, out-of-control eating in individuals with BN. Consistent with this notion, we found that negative affect-induced delays in accumulating health information relative to taste information for higher fat foods were associated with more frequent out-of-control eating episodes in BN.

Revised version: Negative affect is a known precursor to over-eating and dietary rule-breaking (Frayn & Knäuper, 2018). Healthy adults, and even rodents, can exhibit dysregulated food consumption and altered food choices after periods of stress or other negative emotions (Dionysopoulou et al., 2021; Jacques et al., 2019; Konttinen et al., 2010; Tomiyama, 2019). Negative affect’s influence on eating behavior is particularly pronounced among individuals with clinical diagnoses like BN, who tend to have large and out-of-control eating episodes in states of strong negative emotion (Cardi et al., 2015; Wonderlich et al., 2022). Prior studies have demonstrated that outside of binge-eating episodes, and in states of lower negative affect, individuals with BN typically engage in restrictive eating behaviors, consuming some low-fat foods and largely avoiding high-fat foods (Alpers & Tuschen-Caffier, 2004; Bjorlie et al., 2022; Elran-Barak et al., 2015; Kales, 1990). We previously found that women with BN demonstrated these restrictive food choices on a food-related decision task, regardless of their affective state (Gianini et al., 2019). However, by leveraging a computational model that captures the decision-making process underlying ultimate choices, we found that negative affect more strongly biased individuals with BN than healthy adults, increasing the amount of time that taste information alone influenced their decisions involving food. This affect-induced delay in the consideration of healthiness versus tastiness may help account for the fact that negative affect is a common precursor to taste-driven overeating in healthy adults, and the fact that it is a particularly strong predictor of taste-driven, out-of-control eating in individuals with BN. Consistent with this notion, we found that negative affect-induced delays in accumulating health information relative to taste information for higher fat foods were associated with more frequent out-of-control eating episodes in BN.

Earlier version: Previous studies have shown that in states of neutral affect, individuals with BN are more likely to consume foods that are low-fat and considered to be healthy (e.g., salads, fruits), as well as use healthiness information to guide their food choices (Alpers & Tuschen-Caffier, 2004; Bjorlie et al., 2022; Elran-Barak et al., 2015; Kales, 1990; Neveu et al., 2018; Schnepper et al., 2021). Our results demonstrate that in a neutral mood state, despite an initial tendency to focus on taste information before healthiness information, the BN group had stronger weights on healthiness information and negative weights on taste, especially for high-fat foods. This negative weight on taste information (i.e., an aversive assignment to the value of taste) likely facilitated the BN group’s restrictive food choice policy compared to healthy controls. Individuals with BN may compensate for their tendency to focus on the tastiness of high-fat foods by negatively reappraising or deemphasizing the relative importance of taste.

Revised version: Previous studies have shown that in states of neutral affect, individuals with BN are more likely to consume foods that are low-fat and considered to be healthy (e.g., salads, fruits), as well as use healthiness information to guide their food choices (Alpers & Tuschen-Caffier, 2004; Bjorlie et al., 2022; Elran-Barak et al., 2015; Kales, 1990; Neveu et al., 2018; Schnepper et al., 2021). Our results demonstrate that in a neutral mood state, despite an initial tendency to focus on taste information before healthiness information, the BN group had stronger weights on healthiness information and lower, and in fact negatively signed, weights on taste, especially for high-fat foods. This negative weight on taste information (i.e., an aversive assignment to the value of taste) may have facilitated the BN group’s restrictive food choice policy compared to healthy controls. Individuals with BN may compensate for their tendency to focus on the tastiness of high-fat foods by negatively reappraising or deemphasizing the relative importance of taste and highly weighting their food choices based on health.

Earlier version: During binge-eating episodes, which often follow increases in negative affect, individuals with BN tend to consume large amounts of high-fat foods that were previously avoided (Alpers & Tuschen-Caffier, 2001; Berg et al., 2013; Collins et al., 2017; Haedt-Matt & Keel, 2011; Hilbert & Tuschen-Caffier, 2007; Smyth et al., 2007). Although previous studies examined food consumption associated with negative emotions in BN (Cardi et al., 2015), discrete food choice behaviors had not been investigated, and the influence of negative affect on specific aspects of the decision-making process were unclear. We showed that this affect-related change in eating behavior may occur because negative affect disrupts the delicate balance between how much weight tastiness and healthiness carry in influencing food choices and the relative speed with which they begin to exert their influence (Barakchian et al., 2021; HajiHosseini & Hutcherson, 2021; Hutcherson & Tusche, 2022; Lim et al., 2018; Maier et al., 2020; Schubert et al., 2021; Sullivan et al., 2015; Sullivan & Huettel, 2021). When in states of high negative affect, both healthy adults and individuals with BN may consider food’s appealing taste for longer, biasing them towards choosing tastier, high-fat foods. If this temporal advantage for tastiness is long enough, and a positive weight on tastiness is strong enough, individuals with BN could become more vulnerable to out-ofcontrol eating because they could decide to eat tastier or high-fat foods before ever considering their perceived healthiness. This notion is consistent with our finding that greater affect-induced delays in processing healthiness information were associated with more frequent subjective binge episodes—episodes of dysregulated eating that may not include a large amount of food, but like objectively large binge-eating episodes, are dominated by carbohydrates and high-fat foods that are highly palatable (Presseller et al., 2023), and are often spurred by negative affect (Alpers & Tuschen-Caffier, 2001; Berg et al., 2013; Haedt-Matt & Keel, 2011; Hilbert & Tuschen-Caffier, 2007; Smyth et al., 2007).

Revised version: During binge-eating episodes, which often follow increases in negative affect, individuals with BN tend to consume large amounts of high-fat foods that were previously avoided (Alpers & Tuschen-Caffier, 2001; Berg et al., 2013; Collins et al., 2017; Haedt-Matt & Keel, 2011; Hilbert & Tuschen-Caffier, 2007; Smyth et al., 2007). Although prior studies examined food consumption associated with negative emotions in BN (Cardi et al., 2015), discrete food choice behaviors had not been investigated, and the influence of negative affect on specific aspects of the decision-making process were unclear. We showed that this affect-related change in eating behavior may occur because negative affect disrupts the delicate balance between how much weight tastiness and healthiness carry in influencing food choices and the relative speed with which they begin to exert their influence (Barakchian et al., 2021; HajiHosseini & Hutcherson, 2021; Hutcherson & Tusche, 2022; Lim et al., 2018; Maier et al., 2020; Schubert et al., 2021; Sullivan et al., 2015; Sullivan & Huettel, 2021). When in states of high negative affect, both healthy adults and individuals with BN may consider food’s appealing taste for longer, biasing them towards choosing tastier, high-fat foods. If this temporal advantage for tastiness is long enough, and a positive weight on tastiness is strong enough, individuals with BN could become more vulnerable to out-of-control eating because they could decide to eat tastier or high-fat foods before ever considering their perceived healthiness. This notion is consistent with our finding that greater affect-induced delays in processing healthiness information were associated with more frequent subjective binge episodes— episodes of dysregulated eating that may not include a large amount of food, but like objectively large binge-eating episodes, are dominated by carbohydrates and high-fat foods that are highly palatable (Presseller et al., 2023), and are often spurred by negative affect (Alpers & Tuschen-Caffier, 2001; Berg et al., 2013; Haedt-Matt & Keel, 2011; Hilbert & Tuschen-Caffier, 2007; Smyth et al., 2007).

Earlier version: Our findings reveal new insights into the dynamics of food-related decision-making in both healthy adults and BN and may help to inform new interventions. When attribute weights are inconsistent with a larger goal, decision-makers can use shifts in the timing of attribute consideration to achieve a desired outcome (Maier et al., 2020). For example, earlier shifts in the relative onset time of healthiness information could help an individual with stronger weights on tastiness adhere to a diet. Perhaps targeting affect-induced shifts in the timing of attribute consideration could help regulate the over- and under-controlled eating seen in binge-fast cycles. For example, past research in individuals without eating disorders indicates that cueing attention towards the healthiness of foods during a choice task both increased the weight on healthiness information and increased its onset time during decision-making (Barakchian et al., 2021; Maier et al., 2020) but see (Sullivan & Huettel, 2021). Several interventions already focus on modifying attentional biases in individuals with obesity and binge eating (Boutelle et al., 2016; Brockmeyer et al., 2019; Stojek et al., 2018). Our results suggest that such interventions could be targeted even more precisely. One way for individuals to avoid emotion-driven binge eating may be to focus attention away from the taste attributes of tastier, high-fat foods during states of high negative affect. Conversely, during states of low negative affect, increased attention on the positive taste attributes of all foods may be instrumental in motivating individuals with BN to engage in less restrictive eating.

Revised version: Our findings reveal new insights into the dynamics of food-related decision-making in both healthy adults and BN and may help to inform new interventions. When attribute weights are inconsistent with a larger goal, decision-makers can use shifts in the timing of attribute consideration to achieve a desired outcome (Amasino et al., 2019; Maier et al., 2020). For example, earlier shifts in the relative onset time of healthiness information could help an individual with stronger weights on tastiness adhere to a diet. Perhaps targeting affect-induced shifts in the timing of attribute consideration could help regulate the over- and under-controlled eating seen in binge-fast cycles. For example, past research in individuals without eating disorders indicates that cueing attention towards the healthiness of foods during a choice task both increased the weight on healthiness information and increased its onset time during decision-making (Barakchian et al., 2021; Maier et al., 2020) but see (Sullivan & Huettel, 2021). Several interventions already focus on modifying attentional biases in individuals with obesity and binge eating (Boutelle et al., 2016; Brockmeyer et al., 2019; Stojek et al., 2018). Our results suggest that such interventions could be targeted even more precisely. One way for individuals to avoid emotion-driven binge eating may be to focus attention away from the taste attributes of tastier, high-fat foods during states of high negative affect. Conversely, during states of low negative affect, increased attention on the positive taste attributes of all foods may be instrumental in motivating individuals with BN to engage in less restrictive eating.

Earlier version: Although we found an aberrant influence of negative affect on the computations that contribute to food choice in BN, in our paradigm, the influence of negative affect was not strong enough to change final choices in either group. Even after the negative affect induction in BN, choice outcomes were ultimately determined by the advantage in weighted evidence for healthiness as opposed to the advantage in relative timing for tastiness. We speculate that this may have been for two reasons. First, the affect induction method may not have been potent enough. In the current study, post-induction increases in negative affect (dHC = 1.34; dBN = 1.19) were smaller than those reported after other induction methods, such as viewing affectively salient images with congruent music (d = 3.33; (Zhang et al., 2014)), and they may not have been powerful enough to effect changes in ultimate choices. In addition, our negative-affect inducing music and writing exercise may have created an emotional experience that inadequately represents the negative affect that typically precedes emotional eating or binge eating. Qualitative studies of patients with BN indicate that interpersonal stressors and related negative emotions (e.g., resentment) often precipitate binge episodes (Bohon et al., 2021; Wasson, 2003). Perhaps larger increases in negative affect intensity or exposure to the specific types of negative affect which typically precede binge eating could produce a large enough asynchrony in onset timing of attributes and changes in attribute weights to change individuals’ choices (Figure 6). Second, participants knew they would be offered a snack-sized amount of food from a randomly selected trial to consume at the end of the task. This quantity of food may not have provided participants with access to the type of binge-eating experience that negative affect would otherwise precipitate. Further work is needed to test whether the same processes identified in the current study also drive binge eating outside the lab (Wonderlich et al., 2022). Future adjustments to incentive compatibility, with designs that allow for increased consumption beyond snack-sized portions and/or guaranteed access to a wider-array of high-fat foods could clarify whether opportunities to overconsume highly palatable food would generate different patterns of food choice behavior on the task.

Revised version: Although we found an aberrant influence of negative affect on the computations that contribute to food choice in BN, in our paradigm, the influence of negative affect was not strong enough to change final choices in either group. Even after the negative affect induction which decreased taste weights for HC and not BN, choice outcomes in BN were ultimately determined by the advantage in weighted evidence for healthiness as opposed to the advantage in relative timing for tastiness. We speculate that this may have been for two reasons. First, the affect induction method may not have been potent enough. In the current study, post-induction increases in negative affect (dHC = 1.34; dBN = 1.19) were smaller than those reported after other induction methods, such as viewing affectively salient images with congruent music (d = 3.33; (Zhang et al., 2014)), and they may not have been powerful enough to effect changes in ultimate choices. Although the Food Choice Task was completed immediately after the affect induction, we do not know how long its effect lasted because mood was not repeatedly assessed throughout the task. This and previous studies using task-based designs have only assessed mood ratings immediately following affect inductions (Rouhani et al., 2025; Werthmann et al., 2014), but future studies could include continuous assessments throughout the experiment to better quantify potential group differences in how induced affect changes persist and/or fluctuate. In addition, our negative-affect inducing music and writing exercise may have created an emotional experience that inadequately represents the negative affect that typically precedes emotional eating or binge eating. Qualitative studies of patients with BN indicate that interpersonal stressors and related negative emotions (e.g., resentment) often precipitate binge episodes (Bohon et al., 2021; Wasson, 2003). Perhaps larger increases in negative affect intensity or exposure to the specific types of negative affect which typically precede binge eating could produce a large enough asynchrony in onset timing of attributes and changes in attribute weights to change individuals’ choices (Figure 6). Second, participants knew they would be offered a snack-sized amount of food from a randomly selected trial to consume at the end of the task. This quantity of food may not have provided participants with access to the type of binge-eating experience that negative affect would otherwise precipitate. Further work is needed to test whether the same processes identified in the current study also drive binge eating outside the lab (Wonderlich et al., 2022). Future adjustments to incentive compatibility, with designs that allow for increased consumption beyond snack-sized portions and/or guaranteed access to a wider-array of high-fat foods could clarify whether opportunities to overconsume highly palatable food would generate different patterns of food choice behavior on the task. Additionally, future research should include continuous assessments of other relevant factors such as cravings (Konova et al., 2018; Wonderlich et al., 2017) to enrich the connection between mood, discrete food choice, and binge eating.

Revised model of affect-induced binge-eating.

Each trajectory is a simulated individual with BN who considers the taste attribute alone (pink shaded sides of the panels) before considering both tastiness and healthiness attributes together (blue shaded sides of the panels). In both panels the solid purple trajectory illustrates a decision involving a low-fat food, where tastiness information has a weak, positive weight and healthiness information has a strong, positive weight. The dashed blue trajectory illustrates a decision involving a high-fat food, where tastiness information has a strong, positive weight and healthiness information has a strong, negative weight. (A) Low Negative affect. During neutral affect, the high-fat food is trending towards the accept threshold until the onset of aversive healthiness information biases the trajectory towards the reject threshold. (B) High Negative affect. During negative affect, the initial delay is shifted and tastiness information is accumulated longer before healthiness information comes online. With a longer delay in the onset of healthiness information, the evidence accumulation for the high-fat food has enough time to reach the accept threshold.

Earlier version: Of note, the modest sample size (n = 46) may limit the generalizability of our findings, warranting replication in larger samples. However, the study’s repeated-measures design benefits from the added statistical power and sensitivity of within-subjects statistical tests, which reduce error variance by controlling for individual differences (Charness et al., 2012; Judd et al., 2017; Westfall et al., 2014). Moreover, we collected a large quantity of data per subject, which enhances the validity of the findings by providing a more precise estimate of the effects of the experimental manipulation, despite the smaller sample size (Baker et al., 2021; Smith & Little, 2018).

Revised version: Of note, the modest sample size (n = 46) composed exclusively of female participants may limit the generalizability of our findings. Replication in larger samples that include male participants is needed. Nonetheless, the study’s repeated-measures design benefits from the added statistical power and sensitivity of within-subjects statistical tests, which reduce error variance by controlling for individual differences (Charness et al., 2012; Judd et al., 2017; Westfall et al., 2014). Moreover, we collected a large quantity of data per subject, which enhances the validity of the findings by providing a more precise estimate of the effects of the experimental manipulation, despite the smaller sample size (Baker et al., 2021; Smith & Little, 2018).

Earlier version: Overall, our results advance our understanding of how changes in affective states may shift individuals with BN away from typically restrictive dietary choices and towards binge eating. Specifically, negative affect can bias individuals with BN to focus on taste information for longer, and the strength of this effect for high-fat foods predicts symptom severity. Our results add to recent data suggesting that computational modeling can detect subtle alterations in decision-making processes in BN that are sensitive to state changes (Berner et al., 2023), and future research should continue examining the potential connections between negative affect and eating-related decision-making using these analytic approaches.

Revised version: Overall, our results advance our understanding of how changes in affective states may shift individuals with BN away from typically restrictive dietary choices and towards binge eating. Specifically, negative affect can bias individuals with BN to focus on taste information for longer, and the strength of this effect for high-fat foods was associated with symptom severity. Our results add to recent data suggesting that computational modeling can detect subtle alterations in decision-making processes in BN that are sensitive to state changes (Berner et al., 2023), and future research should continue examining the potential connections between negative affect and eating-related decision-making using these analytic approaches.

Materials and Methods

Participants

Data were collected from 25 individuals who met DSM-5 diagnostic criteria for BN and 21 healthy controls, group-matched for age and BMI. All participants were female and aged 18-40 years. The sample consisted of 27 participants who identified as Caucasian (58.7%), 7 as Hispanic (15.2%), 6 as Asian (13.0%), 3 as Black (6.5%), and 3 as mixed race (6.5%). For details on recruitment and exclusions, see (Gianini et al., 2019).

Materials and procedures

Participants completed the tasks over two separate study sessions, each lasting approximately four hours. On each study day, participants were instructed to abstain from consuming food or drinks except for a standardized meal consumed two hours preceding the study session.

Participants completed a Food Choice Task (Steinglass et al., 2015) on two separate days, in two counterbalanced states: after a neutral affect induction, and after a negative affect induction. This task is composed of three phases: Healthiness Rating, Tastiness Rating, and Choice. In each phase, participants viewed the same 43 food items, each categorized as either low-fat (< 30% kcals from fat) or high-fat (> 30% kcals from fat). During the Health Rating phase, participants used a 5-point scale to rate the healthiness of each food item from 1 (Unhealthy) to 3 (Neutral) to 5 (Healthy). During the Taste Rating phase, participants used a 5-point scale to rate the tastiness of each food item from 1 (Bad) to 3 (Neutral) to 5 (Good). The order of the Health Rating and Test Rating phases was counterbalanced. After both phases were completed, one food item that was rated Neutral for both health and taste was selected as a Reference Item for use in the Choice phase. If no item was rated Neutral for healthiness and tastiness, then a food item rated 3 on health and 4 or above on taste was selected following previously established procedures (Steinglass et al., 2015). During the Choice phase, participants made choices between the neutral Reference Item and other food items presented on the computer screen. Participants indicated their choice on a 5-point scale from 1 (No – select the Reference Item) to 3 (Indifference) to 5 (Yes – select the shown food item). Participants were informed that one randomly selected trial would be selected for payout at the end of the experiment and would receive the food item they selected on that trial.

The experiment included an affect induction that included a combination of music and autobiographical writing (Werthmann et al., 2014). During the negative affect induction, participants were asked to write about a recent negative experience while listening to “Adagio for Strings” by Samuel Barber for 8 minutes. During the neutral affect induction, participants were asked to write about the route they took to get to the study site while listening to “Dancing with the Sun” by Celia Felix for 8 minutes. The 65-item Profile of Mood States scale was administered before and after the affect inductions to assess changes in affect (McNair et al., 1989).

Participants also completed the following measures: Eating Disorder Examination Questionnaire (EDE-Q; (Fairburn & Beglin, 1994)); Difficulties in Emotion Regulation Scale (DERS; (Gratz & Roemer, 2004)); Urgency, Premeditation, Perseveration, Sensation Seeking, and Positive Urgency Behavior Scale-Negative Urgency subscale (UPPS-P Negative Urgency; (Lynam et al., 2006)); Beck Depression Inventory (BDI; (Beck et al., 1961)); State Anxiety Inventory (STAI; (Spielberger et al., 1983)).

For additional information, please see (Gianini et al., 2019).

Data analysis

Data transformations and exclusions

Following the procedure from (Gianini et al., 2019), participants’ responses from the Choice phase of the FCT were binarized from their 5-point scale. Responses marked 1 or 2 (i.e., “no”) were determined to be choices of the reference item, responses marked 4 or 5 (i.e., “yes”) were determined to be choices of the presented food item, and responses marked 3 (i.e., “indifferent”) were omitted (total n =599; BN n = 292, HC n = 307).

In contrast to the original study, we excluded trials containing outlier responses times (RTs). Outlier response times likely reflect phenomena outside the decision-making processes of interest, including attention lapses or accidental responses (Cousineau & Chartier, 2010; Ratcliff, 1993). We excluded outlier responses using cutoffs of ±3 standard deviations (SD) from the mean (Berger & Kiefer, 2021). For each participant, the mean and standard deviation of their RTs were calculated and trials containing responses larger/smaller than the mean ±3 SDs were excluded. Additionally, we removed trials where the RTs were greater than 10,000 ms or less than 250 ms after the mean ±3 SDs treatment. Using this method, we removed 2.5% of trials (total n =85; BN n = 51, HC n = 34).

Computational modeling

We used a time-varying diffusion decision model (DDM; (Ratcliff & McKoon, 2008) to study the dynamics of the decision-making during the Food Choice Task. This time-varying model introduces a starting time parameter that indicates the delay with which different attributes enter the evidence accumulation process. This model was previously fitted to food choice data (Lombardi & Hare, 2021; Maier et al., 2020) and lottery choice data (Chen et al., 2022) from healthy adults.

Earlier version: The starting time DDM (stDDM) was estimated using a hierarchical, Bayesian framework implemented with the R package Rjags, which uses the JAGS MCMC sampling algorithm (Plummer, 2003). Our fitting scripts were adapted from those developed by Hsiang-Yu Chen, Gaia Lombardo, and Todd Hare (Chen et al., 2022). This fitting method simultaneously estimates both group- and individual-level parameters. Parameter estimates were fitted to data that included individual-level response times, choices, z-scored health and taste ratings, and Affect Condition. We specified a model where individual parameters were drawn from four separate group-level distributions that varied both by Group (BN or HC) and by Affect Induction (neutral or negative). Within this model, attribute weight parameters (ωTaste and ωHealth) and the relative-starting time parameter (τs) were estimated separately based on food-types (low-fat or high-fat).

Revised version: The starting time DDM (stDDM) was estimated using a hierarchical, Bayesian framework implemented with the R package Rjags, which uses the JAGS MCMC sampling algorithm (Plummer, 2003). Our fitting scripts were adapted from those developed by Hsiang-Yu Chen, Gaia Lombardo, and Todd Hare (Chen et al., 2022). This fitting method simultaneously estimates both group- and individual-level parameters, which improves the reliability of parameters estimated from data with low trial counts per participant (Ratcliff & Childers, 2015; Wiecki et al., 2013). Parameter estimates were fitted to data that included individual-level response times, choices, z-scored health and taste ratings, and Affect Condition. We specified a model where individual parameters were drawn from four separate group-level distributions that varied both by Group (BN or HC) and by Affect Induction (neutral or negative). Within this model, attribute weight parameters (ωTaste and ωHealth) and the relative-starting time parameter (τS) were estimated separately based on food-types (low-fat or high-fat). Using model comparisons and posterior predictive checks, we confirmed that this model best fit our data compared to simpler models (see Supplementary Materials).

Earlier version: Following previously described protocols (Chen et al., 2022), we used group-level priors for attribute weight parameters (ωTaste and ωHealth), the relative-starting time parameter (τs), and the affect induction parameters that were drawn from Gaussian distributions with mean = 0 and standard deviation = 1. These parameter values were then divided by their standard deviations before fitting the model. The priors for boundary separation (α) and non-decision time (τd) were drawn from uniform distributions with ranges 1.0 × 10−4 to 5 and 0 to 10, respectively. The priors for starting point bias (z) were drawn from a beta distribution where both the shape and scale parameters were set to 2. All individual-level priors were drawn from gamma distributions with shape parameters of 1 and scale parameters of 0.1. Posterior estimates were drawn from three chains, each with 100,000 samples (85,000 discarded as burn-in) and thinning every 10 samples. For parametric tests involving parameter estimates, we used the medians of the individual-level posterior distributions. See the Supplementary Materials for model information on model comparisons, model identifiability, posterior predictive checks, and parameter recovery.

Revised version: Following previously described protocols (Chen et al., 2022), we used group-level priors for attribute weight parameters (ωTaste and ωHealth), the relative-starting time parameter τs), and the affect induction parameters that were drawn from Gaussian distributions with mean = 0 and standard deviation = 1. These parameter values were then divided by their standard deviations before fitting the model. The priors for boundary separation (α) and non-decision time (τND) were drawn from uniform distributions with ranges 1.0 × 10−4 to 5 and 0 to 10, respectively. The priors for starting point bias (z) were drawn from a beta distribution where both the shape and scale parameters were set to 2. All individual-level priors were drawn from gamma distributions with shape parameters of 1 and scale parameters of 0.1. Posterior estimates were drawn from three chains, each with 100,000 samples (85,000 discarded as burn-in) and thinning every 10 samples. Convergence among chains was assessed using the Gelman-Rubin statistic (Gelman & Rubin, 1992) using a threshold of 1.1. Model fit was quantified with the Widely Applicable Information Criterion (WAIC) (Watanabe, 2010), which penalizes for model complexity based on the number of parameters to prevent overfitting. For parametric tests involving parameter estimates, we used the medians of the individual-level posterior distributions. See the Supplementary Materials for model information on model comparisons, model identifiability, posterior predictive checks, and parameter recovery.

Earlier version: Multilevel linear regressions were run to evaluate our hypotheses regarding the influence of Food Type, affect induction, BN diagnosis, and their interactions on parameter estimates. Each model included Food Type (coded -1/1 for low-fat/high-fat foods), Affect Condition (coded -1/1 for Neutral/Negative), and Group (coded -1/1 for HC/BN), and their interactions as independent variables. For these models, the intercepts were treated as random effects and statistical significance was evaluated using α = 0.050.

Revised version: Multilevel linear regressions were run to evaluate our hypotheses regarding the influence of Food Type, affect induction, BN diagnosis, and their interactions on parameter estimates. These models were implemented using the lme4 (Bates et al., 2015) and lmertest (Kuznetsova et al., 2013) R packages. Each regression model included Food Type (coded −1/1 for low-fat/high-fat foods), Affect Condition (coded −1/1 for Neutral/Negative), and Group (coded −1/1 for HC/BN), and their interactions as independent variables. For these models, the intercepts were treated as random effects and statistical significance was evaluated using α = 0.050.

Added in the revised version: To unpack observed interaction effects, we ran exploratory simple effects analyses to separately evaluate the influence of independent variables at specific values of other independent variables (Aiken et al., 1991; Brambor et al., 2006; Jaccard & Turrisi, 2003; Winer et al., 1991).

Added in the revised version: We used the R package emmeans to compute estimated marginal means for specified factor combinations (Lenth, 2025), evaluating statistical significance using α = 0.050.

Removed in the revised version: To unpack observed interaction effects, we ran additional, exploratory models to separately evaluate the influence of Food Type, Affect Condition, and their interaction within each group (HC and BN). These models used the same coding conventions and random effects structure as our omnibus tests. To account for multiple comparisons, we applied a Bonferroni correction control the family-wise error rate (FWE) for tests evaluating a specific parameter, using α = 0.025.

Earlier version: Additional fixed-effects regressions were run to assess the relationship between parameter estimates and symptom severity. For these exploratory analyses, we used negative binomial models to predict the threemonth frequency of self-report Objective Binge Episodes (OBE) and Subjective Binge Episodes (SBE). For each of these two models, parameter estimates were entered along with dummy variables for Food Type and Affect Condition, and their interactions, as independent variables. We applied a Bonferroni correction to control for the FWE of these exploratory analyses, using α = 0.025.

Revised version: Additional fixed-effects regressions were run to assess the relationship between parameter estimates and symptom severity. For these exploratory analyses, we used negative binomial models to predict the three- month frequency of self-report Objective Binge Episodes (OBE) and Subjective Binge Episodes (SBE). We implemented these models using the R package glmmTMB (Brooks et al., 2017), with a zero-inflated term included for the model predicting SBE. For each of these two models, individual-level parameter estimates from each condition and food-type were included as independent variables. We applied a Bonferroni correction to control for the FWE of these exploratory analyses, using α = 0.025.